#### Homework-5
##### Total number of points: 70
#### Due date: October 20th, 2022

Before you submit this homework, make sure everything runs as expected. First, restart the kernel (in the menu, select Kernel → Restart) and then run all cells (in the menubar, select Cell → Run All). You can discuss with others regarding the homework but all work must be your own.

This homework will test your knowledge on basics of Python. The Python notebooks shared will be helpful to solve these problems. 

Steps to evaluate your solutions:

Step-1: Ensure you have installed Anaconda (Windows: https://docs.anaconda.com/anaconda/install/windows/ ; Mac:https://docs.anaconda.com/anaconda/install/mac-os/ ; Linux: https://docs.anaconda.com/anaconda/install/linux/)

Step-2: Open the Jupyter Notebook by first launching the anaconda software console

Step-3: Open the .ipynb file and write your solutions at the appropriate location "# YOUR CODE HERE"

Step-4: You can restart the kernel and click run all (in the menubar, select Cell → Run All) on the center-right on the top of this window.

Step-5: Now go to "File" then click on "Download as" then click on "Notebook (.ipynb)" Please DO NOT change the file name and just keep it as .ipynb file format

Step-6: Go to lms.rpi.edu and upload your homework at the appropriate link to submit this homework.

#### Please note that for any question in this assignment you will receive points ONLY if your solution passes all the test cases including hidden testcases as well. So please make sure you try to think all possible scenarios before submitting your answers.  
- Note that hidden tests are present to ensure you are not hardcoding. 
- If caught cheating: 
    - you will receive a score of 0 for the 1st violation. 
    - for repeated incidents, you will receive an automatic 'F' grade and will be reported to the dean of Lally School of Management. 

#### Q1 [10 points]. Please make sure this is correct as the following questions are dependent on this. 
- Load the boston data from sklearn library (library is already provided here below). 
##### Part-1:
- The given target labels (boston.target) are continuous -- convert them into discrete values -- 1 and 2 using this approach and save them in an array `y_true`.
- all the values (v1) in boston.target should become 1 if v1>=5 and v1<23; 
- all the values (v2) in boston.target should become 2 if v2>=23 and v2<51;
- Note that these new/transformed labels 1 and 2 are **integers**
- Please save these new discrete values in the array `y_true`

##### Part-2: 
- Create a list `y_pred` which is of the same length as `y_true` and insert all values as 1.
- Now use `y_true` and `y_pred` to calculate: 1) accuracy and assign it to variable `tacc`; 2) precision and assign it to variable `tprec`; 3) recall and assign it to variable `trecall`.
- Create another list `y_pred_prob` which is of the same length as `y_true` and insert all values as 0.75. 

<!-- grade_id: q1-sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_boston

#Load the dataset and save it as a dataframe
boston = load_boston()
df = pd.DataFrame(boston.data, columns=boston.feature_names)

# Assigning y_true as a list
y_true = []

# Assigning all the v1 and v2 values and inserting them into y_true
bins = np.array([5.0, 23.0, 51.0])
y_true = np.digitize(boston.target, bins)

# Assigning 506 1s into y_pred
y_pred = [1] * len(y_true)

# Calculating accuracy, precision, recall, and roc/auc
from sklearn import metrics
tacc = metrics.accuracy_score(y_true, y_pred)
tprec = metrics.precision_score(y_true, y_pred)
trecall = metrics.recall_score(y_true, y_pred)
tauc = metrics.roc_auc_score(y_true, y_pred)

# Assigning 506 0.75s into y_pred_prob
y_pred_prob = [0.75] * len(y_true)
```

Output:
```
/usr/local/lib/python3.7/dist-packages/sklearn/utils/deprecation.py:87: FutureWarning: Function load_boston is deprecated; `load_boston` is deprecated in 1.0 and will be removed in 1.2.

    The Boston housing prices dataset has an ethical problem. You can refer to
    the documentation of this function for further details.

    The scikit-learn maintainers therefore strongly discourage the use of this
    dataset unless the purpose of the code is to study and educate about
    ethical issues in data science and machine learning.

    In this special case, you can fetch the dataset from the original
    source::

        import pandas as pd
        import numpy as np


        data_url = "http://lib.stat.cmu.edu/datasets/boston"
        raw_df = pd.read_csv(data_url, sep="\s+", skiprows=22, header=None)
        data = np.hstack([raw_df.values[::2, :], raw_df.values[1::2, :2]])
        target = raw_df.values[1::2, 2]

    Alternative datasets include the California housing dataset (i.e.
```

<!-- grade_id: 1, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-1 
#DO NOT MODIFY/DELETE THIS CELL 
assert set(y_true)=={1,2}
assert (list(y_true)).count(1)==312
assert (list(y_true)).count(0)==0
assert (list(y_true)).count(2)==194
```

<!-- grade_id: 2, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-2
#DO NOT MODIFY/DELETE THIS CELL 
assert round(tacc, 2)==0.62
assert round(tprec, 2)==0.62
assert round(trecall, 2)==1.0
assert round(tauc, 2)==0.5
```

#### Q2 [5 points]. Now split the data into training (80%) and testing data (20%) by using these variable names 
- `X_train`: Training feature columns
- `X_test`: Testing feature columns
- `y_train`: Training labels
- `y_test`: Testing labels
- with only parameters df, y_true and 'test_size'. `df` and `y_true` are initialized in the previous question. 

<!-- grade_id: q2-sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
from sklearn.model_selection import train_test_split

X_train = pd.DataFrame()
y_train = pd.DataFrame()
X_test = pd.DataFrame()
y_test = pd.DataFrame()

# Splitting the data into 80% training and 20% testing 
X_train, X_test, y_train, y_test = train_test_split(df, y_true, test_size = 0.2)
```

<!-- grade_id: 3, points: 10, grade: true, solution: false -->
```python
#[10 points] Test cell-3
#DO NOT MODIFY/DELETE THIS CELL 
assert len(X_train)==404
assert len(y_test)==102
assert len(y_train)==404
assert len(X_test)==102
```

#### Q3 [5 points]. Use the `df` in Q1 to perform these operations: 
<ol>
<li> Standardize the data using the StandardScaler() function from sklearn library</li>
<li> Compute principal components using the fit_transform operation on the original dataset </li>
<li> Now find the number of principal components `n_components` that will retain not more than 75% of the information present in the original dataset </li>
</ol>

- Make sure you declare appropriate Python packages which are not provided by default
- Hint -- Use explained_variance_ratio_.cumsum() function we discussed in the class


<!-- grade_id: q3-sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_boston
from sklearn.preprocessing import StandardScaler
from sklearn.decomposition import PCA
import matplotlib.pyplot as plt

#Load the dataset and save it as a dataframe
boston = load_boston()
df = pd.DataFrame(boston.data, columns=boston.feature_names)

n_components=0

# Standardizing data
scaler = StandardScaler()
scaled_data = scaler.fit_transform(df)

# Computing PCA
pca = PCA()
scaled_data_pca = pca.fit_transform(scaled_data)

# Finding the number of principal components 
scaled_data_pca_75 = pca.fit(scaled_data)
pca_cumsum = scaled_data_pca_75.explained_variance_ratio_.cumsum()
n_components = len(np.where(pca_cumsum <= 0.75)[0].tolist())
```

Output:
```
/usr/local/lib/python3.7/dist-packages/sklearn/utils/deprecation.py:87: FutureWarning: Function load_boston is deprecated; `load_boston` is deprecated in 1.0 and will be removed in 1.2.

    The Boston housing prices dataset has an ethical problem. You can refer to
    the documentation of this function for further details.

    The scikit-learn maintainers therefore strongly discourage the use of this
    dataset unless the purpose of the code is to study and educate about
    ethical issues in data science and machine learning.

    In this special case, you can fetch the dataset from the original
    source::

        import pandas as pd
        import numpy as np


        data_url = "http://lib.stat.cmu.edu/datasets/boston"
        raw_df = pd.read_csv(data_url, sep="\s+", skiprows=22, header=None)
        data = np.hstack([raw_df.values[::2, :], raw_df.values[1::2, :2]])
        target = raw_df.values[1::2, 2]

    Alternative datasets include the California housing dataset (i.e.
```

<!-- grade_id: 4, points: 2.5, grade: true, solution: false -->
```python
#[2.5 points] Test cell-4
#DO NOT MODIFY/DELETE THIS CELL 
assert n_components!=0
```

<!-- grade_id: 5, points: 2.5, grade: true, solution: false -->
```python
#[2.5 points] Test cell-5 Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
```

#### Q4 [10 points]. Given 3 different series `S1`, `S2`, `S3`
<ol>
<li>First, using these 3 series build a dataframe 'df' where the column name associated with each series is same as the name of the series such as `S1`, `S2` or `S3`.</li>
<li>Please note that 'df' doesn't contain any class label column. We are building a simple dataframe with 3 columns from 3 different series each considered as a feature.</li>
<li>If you perform PCA operation, how many maximum principal components will you obtain?</li>
<li>No need to do pca but manually assign the value to an integer variable `max_pcs`</li>
</ol>

<!-- grade_id: q4-sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
from sklearn.preprocessing import MinMaxScaler

S1 = pd.Series([1,2,3,4]) #corresponding column name should be 'S1'
S2 = pd.Series([10, 20, 30, 40]) #corresponding column name should be 'S2'
S3 = pd.Series([2,4,6,8]) #corresponding column name should be 'S3'

max_pcs=0 #Update this variable after building the dataframe #No need to perform PCA 

# Creating the dataframe
df = pd.DataFrame({'S1': S1, 'S2': S2, 'S3': S3})

# Scaling data
scaler = MinMaxScaler()
df_rescaled = scaler.fit_transform(df)

pca = PCA()
pca.fit_transform(df_rescaled)

# Assigning maximum principal components into max_pcs
max_pcs = pca.n_components_
```

<!-- grade_id: 6, points: 7.5, grade: true, solution: false -->
```python
#[7.5 points] Test cell-6
#DO NOT MODIFY/DELETE THIS CELL 
assert set(df.columns)=={'S1', 'S2', 'S3'}
assert df['S3'].mean()==5.0
assert df['S1'].mean()==2.5
```

<!-- grade_id: 7, points: 2.5, grade: true, solution: false -->
```python
#[2.5 points] Test cell-7 Hidden tests
#DO NOT MODIFY/DELETE THIS CELL 
```

#### Q5 [10 points]. Linear regression: Use the boston dataset loaded into the dataframe `df`
##### 1. **Without using train_test_split() function**
- use the 1st 400 rows (i.e., index 0,1,....399) where the corresponding features are loaded as `train_X` and corresponding labels as `train_y`
- Use the remaining rows as testing data -- `test_X` and `test_y` for the features and labels respectively.
- Fit a linear regression line using the training data; then use it to predict labels for testing data as shown in the lecture notebook. Please use the default parameters when calling the linear regression function. 

##### 2. Measure the mean-squared error (MSE) 'mse_split1' using the predicted labels with testing labels.
- Round the `mse_split1` to 2 values after the decimal point. 
- Hint -- Check out the libraries given to guess which functions to use to compute MSE.


<!-- grade_id: q5-sol, points: 0, grade: false, solution: true -->
```python

import pandas as pd
import numpy as np
from sklearn.datasets import load_boston
from sklearn import datasets
from sklearn import linear_model
from sklearn.metrics import mean_squared_error


#Load the dataset without class labels and save it as a dataframe
boston = load_boston()
df = pd.DataFrame(boston.data, columns=boston.feature_names)
#class label -- boston.target

train_X = pd.DataFrame()
train_y = pd.DataFrame()
test_X = pd.DataFrame()
test_y = pd.DataFrame()

mse_split1=0

# Assigning train_X and train_y
train_X = df.iloc[:400, :]
train_y = boston.target[:400]

# Assigning test_X and test_y
test_X = df.iloc[400:, :]
test_y = boston.target[400:]

# Linear regression
data_regression = linear_model.LinearRegression()
data_regression.fit(train_X, train_y)
pred = data_regression.predict(test_X)

# Measuring MSE
mse_split1 = round(mean_squared_error(test_y, pred), 2)
```

Output:
```
/usr/local/lib/python3.7/dist-packages/sklearn/utils/deprecation.py:87: FutureWarning: Function load_boston is deprecated; `load_boston` is deprecated in 1.0 and will be removed in 1.2.

    The Boston housing prices dataset has an ethical problem. You can refer to
    the documentation of this function for further details.

    The scikit-learn maintainers therefore strongly discourage the use of this
    dataset unless the purpose of the code is to study and educate about
    ethical issues in data science and machine learning.

    In this special case, you can fetch the dataset from the original
    source::

        import pandas as pd
        import numpy as np


        data_url = "http://lib.stat.cmu.edu/datasets/boston"
        raw_df = pd.read_csv(data_url, sep="\s+", skiprows=22, header=None)
        data = np.hstack([raw_df.values[::2, :], raw_df.values[1::2, :2]])
        target = raw_df.values[1::2, 2]

    Alternative datasets include the California housing dataset (i.e.
```

<!-- grade_id: 8, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-8
#DO NOT MODIFY/DELETE THIS CELL 
assert len(train_X)==400
assert len(test_y)==106
assert mse_split1==37.89

```

<!-- grade_id: 9, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-9 Hidden tests 
#DO NOT MODIFY/DELETE THIS CELL 
```

#### Q6 [10 points]. Repeat the above exercise (Q5) but with different training and testing splits using the `boston` dataset.

###### 1. **Without using train_test_split() function** 
- use the last set of rows from index 253 to the end (i.e., index 253,254,....505) where, the corresponding features are loaded as `train_X` and corresponding labels as `train_y`
- Use the remaining rows (from index 0, 1, 2, ...., 252) as testing data -- `test_X` and `test_y` for the features and labels respectively.
- Fit a linear regression line using the training data; then use it for predicting testing data as shown in the lecture notebook. Please use the default parameters when calling the linear regression function. 

###### 2. Measure the mean-squared error (MSE) 'mse_split2' using the predicted labels with testing labels.
- Round the `mse_split2` to 2 values after the decimal point. 
- Hint -- Check out the libraries given to guess which functions to use to compute MSE.


<!-- grade_id: q6-sol, points: 0, grade: false, solution: true -->
```python

import pandas as pd
import numpy as np
from sklearn.datasets import load_boston
from sklearn import datasets
from sklearn import linear_model
from sklearn.metrics import mean_squared_error

#Load the dataset without class labels and save it as a dataframe
boston = load_boston()
df = pd.DataFrame(boston.data, columns=boston.feature_names)
#class label -- boston.target

train_X = pd.DataFrame()
train_y = pd.DataFrame()
test_X = pd.DataFrame()
test_y = pd.DataFrame()

mse_split2=0

# Assigning train_X and train_y
train_X = df.iloc[253:, :]
train_y = boston.target[253:]

# Assigning test_X and test_y
test_X = df.iloc[:253, :]
test_y = boston.target[:253]

# Linear regression
data_regression2 = linear_model.LinearRegression()
data_regression2.fit(train_X, train_y)
pred2 = data_regression2.predict(test_X)

# Measuring MSE
mse_split2 = round(mean_squared_error(test_y, pred2), 2)
```

Output:
```
/usr/local/lib/python3.7/dist-packages/sklearn/utils/deprecation.py:87: FutureWarning: Function load_boston is deprecated; `load_boston` is deprecated in 1.0 and will be removed in 1.2.

    The Boston housing prices dataset has an ethical problem. You can refer to
    the documentation of this function for further details.

    The scikit-learn maintainers therefore strongly discourage the use of this
    dataset unless the purpose of the code is to study and educate about
    ethical issues in data science and machine learning.

    In this special case, you can fetch the dataset from the original
    source::

        import pandas as pd
        import numpy as np


        data_url = "http://lib.stat.cmu.edu/datasets/boston"
        raw_df = pd.read_csv(data_url, sep="\s+", skiprows=22, header=None)
        data = np.hstack([raw_df.values[::2, :], raw_df.values[1::2, :2]])
        target = raw_df.values[1::2, 2]

    Alternative datasets include the California housing dataset (i.e.
```

<!-- grade_id: 10, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-10
#DO NOT MODIFY/DELETE THIS CELL 
assert len(train_X)==253
assert len(test_y)==253
assert mse_split2==27.22

```

<!-- grade_id: 11, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-11 Hidden tests 
#DO NOT MODIFY/DELETE THIS CELL 
```

#### Q7 [15 points]. We have loaded the boston dataset -- use the dataframe `df` to do: 
###### 1. Convert the target labels -- boston.target from continuous to discrete values -- 1, 2 and 3 using this approach.
- all the values (lets say, each value is represented by v1) in boston.target should become 1 if v1>=5 and v1<20; 
- all the values (lets say, each value is represented by v2) in boston.target should become 2 if v2>=20 and v2<35;
- all the values (lets say, each value is represented by v3) in boston.target should become 3 if v3>=35 and v3<51;

###### 2. **Without using train_test_split()** use the 1st 400 rows (i.e., index 0,1,....399) where, 
- the corresponding features are loaded as `train_X` and corresponding labels as `train_y`
- Use the rest of the rows for testing data -- `test_X` and `test_y` for the features and groundtruth labels respectively.
- Fit a logistic regression with solver=`newton-cg`, C=`1e5`, multi_class=`multinomial`; then use it for predicting testing data as shown in the lecture notebook. If you encounter this warning, please ignore it -- "ConvergenceWarning: newton-cg failed to converge." 

###### 3. Measure the accuracy `acc_split3` using the predicted labels with groundtruth test labels `test_y`.
- Round the `acc_split3` to 2 values after the decimal point. 
- Hint -- Check out the libraries given to guess which functions to use to compute accuracy.

<!-- grade_id: q7-sol, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_boston
from sklearn import datasets
from sklearn import linear_model
from sklearn.metrics import accuracy_score

#Load the dataset without class labels and save it as a dataframe
boston = load_boston()
df = pd.DataFrame(boston.data, columns=boston.feature_names)
#class label -- boston.target

train_X = pd.DataFrame()
train_y = pd.DataFrame()
test_X = pd.DataFrame()
test_y = pd.DataFrame()

acc_split3=0

# Assigning all the v1, v2, and v3 values
bins = np.array([5.0, 20.0, 35.0, 51.0])
target_labels = np.digitize(boston.target, bins)

# Assigning train_X and train_y
train_X = df.iloc[:400, :]
train_y = target_labels[:400]

# Assigning test_X and test_y
test_X = df.iloc[400:, :]
test_y = target_labels[400:]

# Logistic regression
data_regression3 = linear_model.LogisticRegression(solver = 'newton-cg', C = 1e5, multi_class = 'multinomial')
data_regression3.fit(train_X, train_y)
pred2 = data_regression3.predict(test_X)

# Measuring accuracy
acc_split3 = round(accuracy_score(test_y, pred2), 2)
```

Output:
```
/usr/local/lib/python3.7/dist-packages/sklearn/utils/deprecation.py:87: FutureWarning: Function load_boston is deprecated; `load_boston` is deprecated in 1.0 and will be removed in 1.2.

    The Boston housing prices dataset has an ethical problem. You can refer to
    the documentation of this function for further details.

    The scikit-learn maintainers therefore strongly discourage the use of this
    dataset unless the purpose of the code is to study and educate about
    ethical issues in data science and machine learning.

    In this special case, you can fetch the dataset from the original
    source::

        import pandas as pd
        import numpy as np


        data_url = "http://lib.stat.cmu.edu/datasets/boston"
        raw_df = pd.read_csv(data_url, sep="\s+", skiprows=22, header=None)
        data = np.hstack([raw_df.values[::2, :], raw_df.values[1::2, :2]])
        target = raw_df.values[1::2, 2]

    Alternative datasets include the California housing dataset (i.e.
```

<!-- grade_id: 12, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-12
#DO NOT MODIFY/DELETE THIS CELL 
assert len(train_X)==400
assert set(train_y)=={1,2,3}
```

<!-- grade_id: 13, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-13
#DO NOT MODIFY/DELETE THIS CELL 
assert acc_split3==0.74
assert len(test_y)==106
```

<!-- grade_id: 14, points: 5, grade: true, solution: false -->
```python
#[5 points] Test cell-14 Hidden tests 
#DO NOT MODIFY/DELETE THIS CELL 
```